# Part 1: Gradients example

This notebook is designed to illustrate the concept of the gradient.

Let $x \equiv \left[\begin{matrix} x_0 \\ x_1 \end{matrix}\right]$ be a two-dimensional vector, and let
$$
  f(x) \, \equiv \, x^T x \, = \, x_0^2 + x_1^2.
$$

**Exercise 0** (1 point). Implement the Python function, `f(x0, x1)`, so that it computes $f(x)$.

<!-- grade_id: f, points: 0, grade: false, solution: true -->
```python
def f(x0, x1):
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: f_test, points: 1, grade: true, solution: false -->
```python
# Test cell: `f_test`

from numpy import sin, cos, vectorize, isclose
from numpy.random import randn

f_vec = vectorize(f)
theta = randn(1000)
assert all(isclose(f_vec(sin(theta), cos(theta)), 1.0))

print("\n(Passed!)")
```

## The gradient

Let's create a mesh of $[x_0, x_1]$ coordinate values:

```python
from numpy import linspace, meshgrid
x0 = linspace(-2, 2, 11)
x1 = linspace(-2, 2, 11)
X0, X1 = meshgrid(x0, x1)
```

```python
print("X0:\n", X0)
print("X1:\n", X1)
```

Apply `f()` to each of the points in the mesh:

```python
Z = f_vec(X0, X1)
print("Z:\n", Z)
```

Plot `Z[:, :]` as a three-dimensional surface:

```python
from mpl_toolkits.mplot3d import Axes3D
from matplotlib.pyplot import figure, xlabel, ylabel
%matplotlib inline

fig = figure(figsize=(10, 5))

ax3d = fig.add_subplot(121, projection='3d')
ax3d.plot_wireframe(X0, X1, Z)
xlabel('x0')
ylabel('x1')

ax2d = fig.add_subplot(122)
cp = ax2d.contour(X0, X1, Z)
xlabel('x0')
ylabel('x1')
```

The gradient of $f(x)$ with respect to $x$ is
$$
  \nabla_{\!x} \, f(x)
  \, \equiv \,
    \left[\begin{matrix}
       \dfrac{\partial f}{\partial x_0} \\
       \dfrac{\partial f}{\partial x_1}
    \end{matrix}\right]
  \, = \,
    \left[\begin{matrix}
      \dfrac{\partial}{\partial x_0} \left( x_0^2 + x_1^2 \right) \\
      \dfrac{\partial}{\partial x_1} \left( x_0^2 + x_1^2 \right)
    \end{matrix}\right]
  \, = \,
    \left[\begin{matrix}
      2 x_0 \\
      2 x_1
    \end{matrix}\right].
$$

**Exercise 1** (1 point). Implement a function, `grad_f(x0, x1)`, that implements the gradient $\nabla_{\!x} \, f(x)$ shown above. It should return a pair of values since the gradient for this $f(x)$ has two components.

<!-- grade_id: grad_f, points: 0, grade: false, solution: true -->
```python
def grad_f(x0, x1):
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: grad_f_test, points: 1, grade: true, solution: false -->
```python
# Test cell: `grad_f_test`

grad_f_vec = vectorize(grad_f)
z = randn(5)
gx, gy = grad_f_vec(z, -z)
assert all(isclose(gx*0.5, z)) and all(isclose(gy*(-0.5), z)), "Your function might have a bug..."

print("\n(Passed!)")
```

## Visualizing the gradient

Let's generate and plot $\nabla_{\!x} \, f(x)$:

```python
dX0, dX1 = grad_f(X0, X1)

from matplotlib.pyplot import contour, quiver, axis
cp = contour (X0, X1, Z)
quiver(X0, X1, dX0, dX1, scale=40, headwidth=5)
xlabel('x0')
ylabel('x1')
axis('square')
```

**Fin!** If you've gotten this far without errors, your notebook is ready to submit.